<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="category">
   <category><text>$course$/top/Python Sinav/Asama 4 - Fonksiyonlar ve string işlemleri</text></category>
  </question>
  <question type="coderunner">
   <name><text>s4_palindrom - Palindrom</text></name>
   <questiontext format="html"><text><![CDATA[<p><code>palindrom_mu(s)</code> fonksiyonu; büyük/küçük harf ve <b>boşlukları yok sayarak</b> metin palindrom ise <code>True</code>, değilse <code>False</code> döndürsün.</p>

<p>Örnek: <code>&quot;Kay ak&quot;</code> → <code>True</code>.</p>]]></text></questiontext>
   <generalfeedback format="html"><text></text></generalfeedback>
   <defaultgrade>10</defaultgrade>
   <penalty>0.1</penalty>
   <hidden>0</hidden>
   <coderunnertype>python3</coderunnertype>
   <prototypetype>0</prototypetype>
   <allornothing>0</allornothing>
   <penaltyregime>10, 20, ...</penaltyregime>
   <precheck>0</precheck>
   <hidecheck>0</hidecheck>
   <showsource>0</showsource>
   <answerboxlines>18</answerboxlines>
   <answerboxcolumns>100</answerboxcolumns>
   <answerpreload></answerpreload>
   <globalextra></globalextra>
   <useace>1</useace>
   <resultcolumns></resultcolumns>
   <template></template>
   <iscombinatortemplate></iscombinatortemplate>
   <allowmultiplestdins></allowmultiplestdins>
   <answer><![CDATA[def palindrom_mu(s):
    t = s.replace(' ', '').lower()
    return t == t[::-1]
]]></answer>
   <validateonsave>1</validateonsave>
   <testsplitterre></testsplitterre>
   <language></language>
   <acelang></acelang>
   <sandbox></sandbox>
   <grader></grader>
   <cputimelimitsecs></cputimelimitsecs>
   <memlimitmb></memlimitmb>
   <sandboxparams></sandboxparams>
   <testcases>
   <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(palindrom_mu("kayak"))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[True]]></text></expected>
    <extra><text></text></extra>
    <display><text>SHOW</text></display>
   </testcase>
   <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(palindrom_mu("Kay ak"))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[True]]></text></expected>
    <extra><text></text></extra>
    <display><text>SHOW</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(palindrom_mu("python"))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[False]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(palindrom_mu(""))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[True]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(palindrom_mu("A man a plan a canal Panama"))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[True]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(palindrom_mu("ab"))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[False]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(palindrom_mu("Aa"))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[True]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   </testcases>
  </question>
  <question type="coderunner">
   <name><text>s4_sesli - Sesli harf sayısı</text></name>
   <questiontext format="html"><text><![CDATA[<p><code>sesli_say(s)</code> fonksiyonu, metindeki sesli harf (<code>aeıioöuü</code>, büyük/küçük fark etmez) sayısını döndürsün.</p>]]></text></questiontext>
   <generalfeedback format="html"><text></text></generalfeedback>
   <defaultgrade>10</defaultgrade>
   <penalty>0.1</penalty>
   <hidden>0</hidden>
   <coderunnertype>python3</coderunnertype>
   <prototypetype>0</prototypetype>
   <allornothing>0</allornothing>
   <penaltyregime>10, 20, ...</penaltyregime>
   <precheck>0</precheck>
   <hidecheck>0</hidecheck>
   <showsource>0</showsource>
   <answerboxlines>18</answerboxlines>
   <answerboxcolumns>100</answerboxcolumns>
   <answerpreload></answerpreload>
   <globalextra></globalextra>
   <useace>1</useace>
   <resultcolumns></resultcolumns>
   <template></template>
   <iscombinatortemplate></iscombinatortemplate>
   <allowmultiplestdins></allowmultiplestdins>
   <answer><![CDATA[def sesli_say(s):
    return sum(1 for c in s.lower() if c in 'aeıioöuü')
]]></answer>
   <validateonsave>1</validateonsave>
   <testsplitterre></testsplitterre>
   <language></language>
   <acelang></acelang>
   <sandbox></sandbox>
   <grader></grader>
   <cputimelimitsecs></cputimelimitsecs>
   <memlimitmb></memlimitmb>
   <sandboxparams></sandboxparams>
   <testcases>
   <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(sesli_say("merhaba"))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[3]]></text></expected>
    <extra><text></text></extra>
    <display><text>SHOW</text></display>
   </testcase>
   <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(sesli_say("PYTHON"))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[1]]></text></expected>
    <extra><text></text></extra>
    <display><text>SHOW</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(sesli_say(""))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[0]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(sesli_say("ÜÖİ"))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[3]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(sesli_say("kıvırcık"))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[3]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(sesli_say("bcdfg"))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[0]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(sesli_say("AEIOU"))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[5]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   </testcases>
  </question>
  <question type="coderunner">
   <name><text>s4_kelime - En uzun kelime</text></name>
   <questiontext format="html"><text><![CDATA[<p>Bir satır metin okuyun (kelimeler boşlukla ayrılmış). Önce kelime sayısını, sonra <b>en uzun kelimeyi</b> ayrı satırlarda yazdırın. Eşitlik varsa <b>ilk</b> geçen kelime seçilir. Satırda en az bir kelime olacaktır.</p>]]></text></questiontext>
   <generalfeedback format="html"><text></text></generalfeedback>
   <defaultgrade>10</defaultgrade>
   <penalty>0.1</penalty>
   <hidden>0</hidden>
   <coderunnertype>python3</coderunnertype>
   <prototypetype>0</prototypetype>
   <allornothing>0</allornothing>
   <penaltyregime>10, 20, ...</penaltyregime>
   <precheck>0</precheck>
   <hidecheck>0</hidecheck>
   <showsource>0</showsource>
   <answerboxlines>18</answerboxlines>
   <answerboxcolumns>100</answerboxcolumns>
   <answerpreload></answerpreload>
   <globalextra></globalextra>
   <useace>1</useace>
   <resultcolumns></resultcolumns>
   <template></template>
   <iscombinatortemplate></iscombinatortemplate>
   <allowmultiplestdins></allowmultiplestdins>
   <answer><![CDATA[kelimeler = input().split()
print(len(kelimeler))
en = ''
for k in kelimeler:
    if len(k) > len(en):
        en = k
print(en)
]]></answer>
   <validateonsave>1</validateonsave>
   <testsplitterre></testsplitterre>
   <language></language>
   <acelang></acelang>
   <sandbox></sandbox>
   <grader></grader>
   <cputimelimitsecs></cputimelimitsecs>
   <memlimitmb></memlimitmb>
   <sandboxparams></sandboxparams>
   <testcases>
   <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[]]></text></testcode>
    <stdin><text><![CDATA[bugun hava cok guzel
]]></text></stdin>
    <expected><text><![CDATA[4
bugun]]></text></expected>
    <extra><text></text></extra>
    <display><text>SHOW</text></display>
   </testcase>
   <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[]]></text></testcode>
    <stdin><text><![CDATA[a bb ccc dd
]]></text></stdin>
    <expected><text><![CDATA[4
ccc]]></text></expected>
    <extra><text></text></extra>
    <display><text>SHOW</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[]]></text></testcode>
    <stdin><text><![CDATA[tek
]]></text></stdin>
    <expected><text><![CDATA[1
tek]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[]]></text></testcode>
    <stdin><text><![CDATA[bir iki uc dort bes
]]></text></stdin>
    <expected><text><![CDATA[5
dort]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[]]></text></testcode>
    <stdin><text><![CDATA[ab cd ef
]]></text></stdin>
    <expected><text><![CDATA[3
ab]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[]]></text></testcode>
    <stdin><text><![CDATA[  cok   bosluk   var  
]]></text></stdin>
    <expected><text><![CDATA[3
bosluk]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   </testcases>
  </question>
  <question type="coderunner">
   <name><text>s4_sezar - Sezar şifresi</text></name>
   <questiontext format="html"><text><![CDATA[<p><code>sifrele(metin, kaydir)</code> fonksiyonunu yazın. Yalnızca <code>a-z</code> ve <code>A-Z</code> harfleri alfabede <code>kaydir</code> kadar ileri kaydırılır (z&#x27;den sonra a&#x27;ya dönülür, büyük/küçük harf korunur). Diğer karakterler <b>aynen</b> kalır. <code>kaydir</code> negatif olabilir.</p>]]></text></questiontext>
   <generalfeedback format="html"><text></text></generalfeedback>
   <defaultgrade>10</defaultgrade>
   <penalty>0.1</penalty>
   <hidden>0</hidden>
   <coderunnertype>python3</coderunnertype>
   <prototypetype>0</prototypetype>
   <allornothing>0</allornothing>
   <penaltyregime>10, 20, ...</penaltyregime>
   <precheck>0</precheck>
   <hidecheck>0</hidecheck>
   <showsource>0</showsource>
   <answerboxlines>18</answerboxlines>
   <answerboxcolumns>100</answerboxcolumns>
   <answerpreload></answerpreload>
   <globalextra></globalextra>
   <useace>1</useace>
   <resultcolumns></resultcolumns>
   <template></template>
   <iscombinatortemplate></iscombinatortemplate>
   <allowmultiplestdins></allowmultiplestdins>
   <answer><![CDATA[def sifrele(metin, kaydir):
    sonuc = ''
    for c in metin:
        if 'a' <= c <= 'z':
            sonuc += chr((ord(c) - 97 + kaydir) % 26 + 97)
        elif 'A' <= c <= 'Z':
            sonuc += chr((ord(c) - 65 + kaydir) % 26 + 65)
        else:
            sonuc += c
    return sonuc
]]></answer>
   <validateonsave>1</validateonsave>
   <testsplitterre></testsplitterre>
   <language></language>
   <acelang></acelang>
   <sandbox></sandbox>
   <grader></grader>
   <cputimelimitsecs></cputimelimitsecs>
   <memlimitmb></memlimitmb>
   <sandboxparams></sandboxparams>
   <testcases>
   <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(sifrele("abc", 1))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[bcd]]></text></expected>
    <extra><text></text></extra>
    <display><text>SHOW</text></display>
   </testcase>
   <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(sifrele("xyz", 3))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[abc]]></text></expected>
    <extra><text></text></extra>
    <display><text>SHOW</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(sifrele("Hello, World!", 5))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[Mjqqt, Btwqi!]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(sifrele("abc", 0))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[abc]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(sifrele("def", -3))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[abc]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(sifrele("Python 3.12", 13))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[Clguba 3.12]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(sifrele("abc", 27))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[bcd]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   </testcases>
  </question>
  <question type="coderunner">
   <name><text>s4_anagram - Anagram</text></name>
   <questiontext format="html"><text><![CDATA[<p><code>anagram_mi(a, b)</code> fonksiyonu iki kelime birbirinin anagramı ise <code>True</code> döndürsün. Büyük/küçük harf ve boşluklar yok sayılır.</p>]]></text></questiontext>
   <generalfeedback format="html"><text></text></generalfeedback>
   <defaultgrade>10</defaultgrade>
   <penalty>0.1</penalty>
   <hidden>0</hidden>
   <coderunnertype>python3</coderunnertype>
   <prototypetype>0</prototypetype>
   <allornothing>0</allornothing>
   <penaltyregime>10, 20, ...</penaltyregime>
   <precheck>0</precheck>
   <hidecheck>0</hidecheck>
   <showsource>0</showsource>
   <answerboxlines>18</answerboxlines>
   <answerboxcolumns>100</answerboxcolumns>
   <answerpreload></answerpreload>
   <globalextra></globalextra>
   <useace>1</useace>
   <resultcolumns></resultcolumns>
   <template></template>
   <iscombinatortemplate></iscombinatortemplate>
   <allowmultiplestdins></allowmultiplestdins>
   <answer><![CDATA[def anagram_mi(a, b):
    f = lambda s: sorted(s.replace(' ', '').lower())
    return f(a) == f(b)
]]></answer>
   <validateonsave>1</validateonsave>
   <testsplitterre></testsplitterre>
   <language></language>
   <acelang></acelang>
   <sandbox></sandbox>
   <grader></grader>
   <cputimelimitsecs></cputimelimitsecs>
   <memlimitmb></memlimitmb>
   <sandboxparams></sandboxparams>
   <testcases>
   <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(anagram_mi("listen", "silent"))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[True]]></text></expected>
    <extra><text></text></extra>
    <display><text>SHOW</text></display>
   </testcase>
   <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(anagram_mi("Dormitory", "dirty room"))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[True]]></text></expected>
    <extra><text></text></extra>
    <display><text>SHOW</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(anagram_mi("abc", "abd"))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[False]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(anagram_mi("aab", "abb"))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[False]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(anagram_mi("", ""))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[True]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(anagram_mi("abc", "ab"))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[False]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   </testcases>
  </question>
</quiz>
