<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="category">
   <category><text>$course$/top/Python Sinav/Asama 5 - Liste, sözlük, küme, matris</text></category>
  </question>
  <question type="coderunner">
   <name><text>s5_ikinci - İkinci en büyük</text></name>
   <questiontext format="html"><text><![CDATA[<p><code>ikinci_buyuk(liste)</code> fonksiyonu listedeki <b>farklı</b> değerler arasında ikinci en büyüğünü döndürsün. Farklı iki değer yoksa <code>None</code> döndürsün.</p>]]></text></questiontext>
   <generalfeedback format="html"><text></text></generalfeedback>
   <defaultgrade>10</defaultgrade>
   <penalty>0.1</penalty>
   <hidden>0</hidden>
   <coderunnertype>python3</coderunnertype>
   <prototypetype>0</prototypetype>
   <allornothing>0</allornothing>
   <penaltyregime>10, 20, ...</penaltyregime>
   <precheck>0</precheck>
   <hidecheck>0</hidecheck>
   <showsource>0</showsource>
   <answerboxlines>18</answerboxlines>
   <answerboxcolumns>100</answerboxcolumns>
   <answerpreload></answerpreload>
   <globalextra></globalextra>
   <useace>1</useace>
   <resultcolumns></resultcolumns>
   <template></template>
   <iscombinatortemplate></iscombinatortemplate>
   <allowmultiplestdins></allowmultiplestdins>
   <answer><![CDATA[def ikinci_buyuk(liste):
    s = sorted(set(liste), reverse=True)
    return s[1] if len(s) >= 2 else None
]]></answer>
   <validateonsave>1</validateonsave>
   <testsplitterre></testsplitterre>
   <language></language>
   <acelang></acelang>
   <sandbox></sandbox>
   <grader></grader>
   <cputimelimitsecs></cputimelimitsecs>
   <memlimitmb></memlimitmb>
   <sandboxparams></sandboxparams>
   <testcases>
   <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(ikinci_buyuk([1, 2, 3]))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[2]]></text></expected>
    <extra><text></text></extra>
    <display><text>SHOW</text></display>
   </testcase>
   <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(ikinci_buyuk([5, 5, 5]))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[None]]></text></expected>
    <extra><text></text></extra>
    <display><text>SHOW</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(ikinci_buyuk([10, 9, 10, 8]))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[9]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(ikinci_buyuk([]))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[None]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(ikinci_buyuk([7]))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[None]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(ikinci_buyuk([-1, -2, -3]))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[-2]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(ikinci_buyuk([2, 1]))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[1]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   </testcases>
  </question>
  <question type="coderunner">
   <name><text>s5_tekrarsiz - Tekrarları sil</text></name>
   <questiontext format="html"><text><![CDATA[<p><code>tekrarsiz(liste)</code> fonksiyonu, tekrar eden elemanları silip <b>ilk geçişlerin sırasını koruyarak</b> yeni bir liste döndürsün. Girdi listesini değiştirmeyin.</p>]]></text></questiontext>
   <generalfeedback format="html"><text></text></generalfeedback>
   <defaultgrade>10</defaultgrade>
   <penalty>0.1</penalty>
   <hidden>0</hidden>
   <coderunnertype>python3</coderunnertype>
   <prototypetype>0</prototypetype>
   <allornothing>0</allornothing>
   <penaltyregime>10, 20, ...</penaltyregime>
   <precheck>0</precheck>
   <hidecheck>0</hidecheck>
   <showsource>0</showsource>
   <answerboxlines>18</answerboxlines>
   <answerboxcolumns>100</answerboxcolumns>
   <answerpreload></answerpreload>
   <globalextra></globalextra>
   <useace>1</useace>
   <resultcolumns></resultcolumns>
   <template></template>
   <iscombinatortemplate></iscombinatortemplate>
   <allowmultiplestdins></allowmultiplestdins>
   <answer><![CDATA[def tekrarsiz(liste):
    gorulen = set()
    sonuc = []
    for x in liste:
        if x not in gorulen:
            gorulen.add(x)
            sonuc.append(x)
    return sonuc
]]></answer>
   <validateonsave>1</validateonsave>
   <testsplitterre></testsplitterre>
   <language></language>
   <acelang></acelang>
   <sandbox></sandbox>
   <grader></grader>
   <cputimelimitsecs></cputimelimitsecs>
   <memlimitmb></memlimitmb>
   <sandboxparams></sandboxparams>
   <testcases>
   <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(tekrarsiz([1, 2, 2, 3, 1]))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[[1, 2, 3]]]></text></expected>
    <extra><text></text></extra>
    <display><text>SHOW</text></display>
   </testcase>
   <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(tekrarsiz([]))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[[]]]></text></expected>
    <extra><text></text></extra>
    <display><text>SHOW</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(tekrarsiz(['a', 'b', 'a']))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[['a', 'b']]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(tekrarsiz([3, 3, 3]))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[[3]]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(tekrarsiz([5, 4, 3]))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[[5, 4, 3]]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[l = [1, 1, 2]
tekrarsiz(l)
print(l)]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[[1, 1, 2]]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   </testcases>
  </question>
  <question type="coderunner">
   <name><text>s5_frekans - Harf frekansı</text></name>
   <questiontext format="html"><text><![CDATA[<p>Bir satır metin okuyun. Harfleri (küçük harfe çevirerek, boşluk ve diğer karakterleri atlayarak) sayın. Her harfi <code>harf: adet</code> biçiminde, <b>önce adedi çok olan, eşitlikte alfabetik</b> sırayla ayrı satırlarda yazdırın.</p>]]></text></questiontext>
   <generalfeedback format="html"><text></text></generalfeedback>
   <defaultgrade>10</defaultgrade>
   <penalty>0.1</penalty>
   <hidden>0</hidden>
   <coderunnertype>python3</coderunnertype>
   <prototypetype>0</prototypetype>
   <allornothing>0</allornothing>
   <penaltyregime>10, 20, ...</penaltyregime>
   <precheck>0</precheck>
   <hidecheck>0</hidecheck>
   <showsource>0</showsource>
   <answerboxlines>18</answerboxlines>
   <answerboxcolumns>100</answerboxcolumns>
   <answerpreload></answerpreload>
   <globalextra></globalextra>
   <useace>1</useace>
   <resultcolumns></resultcolumns>
   <template></template>
   <iscombinatortemplate></iscombinatortemplate>
   <allowmultiplestdins></allowmultiplestdins>
   <answer><![CDATA[s = input().lower()
sayac = {}
for c in s:
    if c.isalpha():
        sayac[c] = sayac.get(c, 0) + 1
for harf, adet in sorted(sayac.items(), key=lambda x: (-x[1], x[0])):
    print(f'{harf}: {adet}')
]]></answer>
   <validateonsave>1</validateonsave>
   <testsplitterre></testsplitterre>
   <language></language>
   <acelang></acelang>
   <sandbox></sandbox>
   <grader></grader>
   <cputimelimitsecs></cputimelimitsecs>
   <memlimitmb></memlimitmb>
   <sandboxparams></sandboxparams>
   <testcases>
   <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[]]></text></testcode>
    <stdin><text><![CDATA[banana
]]></text></stdin>
    <expected><text><![CDATA[a: 3
n: 2
b: 1]]></text></expected>
    <extra><text></text></extra>
    <display><text>SHOW</text></display>
   </testcase>
   <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[]]></text></testcode>
    <stdin><text><![CDATA[Hello World
]]></text></stdin>
    <expected><text><![CDATA[l: 3
o: 2
d: 1
e: 1
h: 1
r: 1
w: 1]]></text></expected>
    <extra><text></text></extra>
    <display><text>SHOW</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[]]></text></testcode>
    <stdin><text><![CDATA[aabbcc
]]></text></stdin>
    <expected><text><![CDATA[a: 2
b: 2
c: 2]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[]]></text></testcode>
    <stdin><text><![CDATA[a
]]></text></stdin>
    <expected><text><![CDATA[a: 1]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[]]></text></testcode>
    <stdin><text><![CDATA[x y z x
]]></text></stdin>
    <expected><text><![CDATA[x: 2
y: 1
z: 1]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[]]></text></testcode>
    <stdin><text><![CDATA[123 !!
]]></text></stdin>
    <expected><text><![CDATA[]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[]]></text></testcode>
    <stdin><text><![CDATA[Python ve Java
]]></text></stdin>
    <expected><text><![CDATA[a: 2
v: 2
e: 1
h: 1
j: 1
n: 1
o: 1
p: 1
t: 1
y: 1]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   </testcases>
  </question>
  <question type="coderunner">
   <name><text>s5_matris - Matris devriği</text></name>
   <questiontext format="html"><text><![CDATA[<p><code>devrik(m)</code> fonksiyonu, iç içe liste ile verilen matrisin devriğini (transpose) yeni bir liste olarak döndürsün. Boş matris için <code>[]</code> döndürün.</p>]]></text></questiontext>
   <generalfeedback format="html"><text></text></generalfeedback>
   <defaultgrade>10</defaultgrade>
   <penalty>0.1</penalty>
   <hidden>0</hidden>
   <coderunnertype>python3</coderunnertype>
   <prototypetype>0</prototypetype>
   <allornothing>0</allornothing>
   <penaltyregime>10, 20, ...</penaltyregime>
   <precheck>0</precheck>
   <hidecheck>0</hidecheck>
   <showsource>0</showsource>
   <answerboxlines>18</answerboxlines>
   <answerboxcolumns>100</answerboxcolumns>
   <answerpreload></answerpreload>
   <globalextra></globalextra>
   <useace>1</useace>
   <resultcolumns></resultcolumns>
   <template></template>
   <iscombinatortemplate></iscombinatortemplate>
   <allowmultiplestdins></allowmultiplestdins>
   <answer><![CDATA[def devrik(m):
    if not m:
        return []
    return [[m[i][j] for i in range(len(m))] for j in range(len(m[0]))]
]]></answer>
   <validateonsave>1</validateonsave>
   <testsplitterre></testsplitterre>
   <language></language>
   <acelang></acelang>
   <sandbox></sandbox>
   <grader></grader>
   <cputimelimitsecs></cputimelimitsecs>
   <memlimitmb></memlimitmb>
   <sandboxparams></sandboxparams>
   <testcases>
   <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(devrik([[1, 2], [3, 4]]))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[[[1, 3], [2, 4]]]]></text></expected>
    <extra><text></text></extra>
    <display><text>SHOW</text></display>
   </testcase>
   <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(devrik([[1, 2, 3], [4, 5, 6]]))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[[[1, 4], [2, 5], [3, 6]]]]></text></expected>
    <extra><text></text></extra>
    <display><text>SHOW</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(devrik([[1], [2], [3]]))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[[[1, 2, 3]]]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(devrik([]))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[[]]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(devrik([[7]]))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[[[7]]]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(devrik([[1, 2, 3]]))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[[[1], [2], [3]]]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   </testcases>
  </question>
  <question type="coderunner">
   <name><text>s5_notlar - Öğrenci not sözlüğü</text></name>
   <questiontext format="html"><text><![CDATA[<p><code>en_basarili(notlar)</code> fonksiyonu <code>{isim: [not, not, ...]}</code> sözlüğü alır. Ortalaması en yüksek öğrencinin ismini döndürsün. Eşitlikte alfabetik olarak <b>ilk</b> isim seçilsin. Sözlük boşsa <code>None</code>. Notu olmayan öğrencinin ortalaması 0 sayılır.</p>]]></text></questiontext>
   <generalfeedback format="html"><text></text></generalfeedback>
   <defaultgrade>10</defaultgrade>
   <penalty>0.1</penalty>
   <hidden>0</hidden>
   <coderunnertype>python3</coderunnertype>
   <prototypetype>0</prototypetype>
   <allornothing>0</allornothing>
   <penaltyregime>10, 20, ...</penaltyregime>
   <precheck>0</precheck>
   <hidecheck>0</hidecheck>
   <showsource>0</showsource>
   <answerboxlines>18</answerboxlines>
   <answerboxcolumns>100</answerboxcolumns>
   <answerpreload></answerpreload>
   <globalextra></globalextra>
   <useace>1</useace>
   <resultcolumns></resultcolumns>
   <template></template>
   <iscombinatortemplate></iscombinatortemplate>
   <allowmultiplestdins></allowmultiplestdins>
   <answer><![CDATA[def en_basarili(notlar):
    en_iyi, en_ort = None, -1
    for isim in sorted(notlar):
        l = notlar[isim]
        ort = sum(l) / len(l) if l else 0
        if ort > en_ort:
            en_iyi, en_ort = isim, ort
    return en_iyi
]]></answer>
   <validateonsave>1</validateonsave>
   <testsplitterre></testsplitterre>
   <language></language>
   <acelang></acelang>
   <sandbox></sandbox>
   <grader></grader>
   <cputimelimitsecs></cputimelimitsecs>
   <memlimitmb></memlimitmb>
   <sandboxparams></sandboxparams>
   <testcases>
   <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(en_basarili({'ali': [50, 60], 'veli': [90, 70], 'ayse': [100, 50]}))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[veli]]></text></expected>
    <extra><text></text></extra>
    <display><text>SHOW</text></display>
   </testcase>
   <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(en_basarili({}))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[None]]></text></expected>
    <extra><text></text></extra>
    <display><text>SHOW</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(en_basarili({'zeynep': [80], 'ali': [80]}))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[ali]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(en_basarili({'a': [], 'b': [1]}))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[b]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(en_basarili({'x': [100, 100, 100]}))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[x]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(en_basarili({'a': [], 'b': []}))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[a]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   </testcases>
  </question>
  <question type="coderunner">
   <name><text>s5_toplam_ikili - İki sayı toplamı (Two Sum)</text></name>
   <questiontext format="html"><text><![CDATA[<p><code>iki_toplam(liste, hedef)</code> fonksiyonu, toplamı <code>hedef</code> eden iki <b>farklı indeksi</b> <code>(i, j)</code> (<code>i &lt; j</code>) demeti olarak döndürsün. Birden çok çözüm varsa <code>i</code> en küçük, sonra <code>j</code> en küçük olan seçilir. Çözüm yoksa <code>None</code>.</p>]]></text></questiontext>
   <generalfeedback format="html"><text></text></generalfeedback>
   <defaultgrade>10</defaultgrade>
   <penalty>0.1</penalty>
   <hidden>0</hidden>
   <coderunnertype>python3</coderunnertype>
   <prototypetype>0</prototypetype>
   <allornothing>0</allornothing>
   <penaltyregime>10, 20, ...</penaltyregime>
   <precheck>0</precheck>
   <hidecheck>0</hidecheck>
   <showsource>0</showsource>
   <answerboxlines>18</answerboxlines>
   <answerboxcolumns>100</answerboxcolumns>
   <answerpreload></answerpreload>
   <globalextra></globalextra>
   <useace>1</useace>
   <resultcolumns></resultcolumns>
   <template></template>
   <iscombinatortemplate></iscombinatortemplate>
   <allowmultiplestdins></allowmultiplestdins>
   <answer><![CDATA[def iki_toplam(liste, hedef):
    for i in range(len(liste)):
        for j in range(i + 1, len(liste)):
            if liste[i] + liste[j] == hedef:
                return (i, j)
    return None
]]></answer>
   <validateonsave>1</validateonsave>
   <testsplitterre></testsplitterre>
   <language></language>
   <acelang></acelang>
   <sandbox></sandbox>
   <grader></grader>
   <cputimelimitsecs></cputimelimitsecs>
   <memlimitmb></memlimitmb>
   <sandboxparams></sandboxparams>
   <testcases>
   <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(iki_toplam([2, 7, 11, 15], 9))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[(0, 1)]]></text></expected>
    <extra><text></text></extra>
    <display><text>SHOW</text></display>
   </testcase>
   <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(iki_toplam([3, 2, 4], 6))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[(1, 2)]]></text></expected>
    <extra><text></text></extra>
    <display><text>SHOW</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(iki_toplam([1, 2, 3], 10))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[None]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(iki_toplam([3, 3], 6))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[(0, 1)]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(iki_toplam([], 0))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[None]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(iki_toplam([5], 10))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[None]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(iki_toplam([1, 4, 5, 6, 4], 8))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[(1, 4)]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(iki_toplam([0, 4, 3, 0], 0))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[(0, 3)]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   </testcases>
  </question>
</quiz>
