<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="category">
   <category><text>$course$/top/Python Sinav/Final - Paket A</text></category>
  </question>
  <question type="coderunner">
   <name><text>s1_ortalama - Not ortalaması [Paket A]</text></name>
   <questiontext format="html"><text><![CDATA[<p>Üç vize notunu (ayrı satırlarda, tam sayı) okuyun. Ortalamalarını <b>iki ondalık basamakla</b> yazdırın.</p>

<p>Örnek: 70, 80, 85 → <code>78.33</code></p>]]></text></questiontext>
   <generalfeedback format="html"><text></text></generalfeedback>
   <defaultgrade>10</defaultgrade>
   <penalty>0.1</penalty>
   <hidden>0</hidden>
   <coderunnertype>python3</coderunnertype>
   <prototypetype>0</prototypetype>
   <allornothing>0</allornothing>
   <penaltyregime>10, 20, ...</penaltyregime>
   <precheck>0</precheck>
   <hidecheck>0</hidecheck>
   <showsource>0</showsource>
   <answerboxlines>18</answerboxlines>
   <answerboxcolumns>100</answerboxcolumns>
   <answerpreload></answerpreload>
   <globalextra></globalextra>
   <useace>1</useace>
   <resultcolumns></resultcolumns>
   <template></template>
   <iscombinatortemplate></iscombinatortemplate>
   <allowmultiplestdins></allowmultiplestdins>
   <answer><![CDATA[a = int(input())
b = int(input())
c = int(input())
print(f'{(a + b + c) / 3:.2f}')
]]></answer>
   <validateonsave>1</validateonsave>
   <testsplitterre></testsplitterre>
   <language></language>
   <acelang></acelang>
   <sandbox></sandbox>
   <grader></grader>
   <cputimelimitsecs></cputimelimitsecs>
   <memlimitmb></memlimitmb>
   <sandboxparams></sandboxparams>
   <testcases>
   <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[]]></text></testcode>
    <stdin><text><![CDATA[70
80
85
]]></text></stdin>
    <expected><text><![CDATA[78.33]]></text></expected>
    <extra><text></text></extra>
    <display><text>SHOW</text></display>
   </testcase>
   <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[]]></text></testcode>
    <stdin><text><![CDATA[100
100
100
]]></text></stdin>
    <expected><text><![CDATA[100.00]]></text></expected>
    <extra><text></text></extra>
    <display><text>SHOW</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[]]></text></testcode>
    <stdin><text><![CDATA[0
0
1
]]></text></stdin>
    <expected><text><![CDATA[0.33]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[]]></text></testcode>
    <stdin><text><![CDATA[55
60
65
]]></text></stdin>
    <expected><text><![CDATA[60.00]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[]]></text></testcode>
    <stdin><text><![CDATA[99
98
97
]]></text></stdin>
    <expected><text><![CDATA[98.00]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   </testcases>
  </question>
  <question type="coderunner">
   <name><text>s2_harf - Harf notu [Paket A]</text></name>
   <questiontext format="html"><text><![CDATA[<p><code>harf(not_)</code> fonksiyonunu yazın:</p>

<ul>
<li>90–100 → <code>AA</code></li>
<li>80–89 → <code>BA</code></li>
<li>70–79 → <code>BB</code></li>
<li>60–69 → <code>CB</code></li>
<li>50–59 → <code>CC</code></li>
<li>0–49 → <code>FF</code></li>
<li>0&#x27;dan küçük veya 100&#x27;den büyük → <code>Gecersiz</code></li>
</ul>]]></text></questiontext>
   <generalfeedback format="html"><text></text></generalfeedback>
   <defaultgrade>15</defaultgrade>
   <penalty>0.1</penalty>
   <hidden>0</hidden>
   <coderunnertype>python3</coderunnertype>
   <prototypetype>0</prototypetype>
   <allornothing>0</allornothing>
   <penaltyregime>10, 20, ...</penaltyregime>
   <precheck>0</precheck>
   <hidecheck>0</hidecheck>
   <showsource>0</showsource>
   <answerboxlines>18</answerboxlines>
   <answerboxcolumns>100</answerboxcolumns>
   <answerpreload></answerpreload>
   <globalextra></globalextra>
   <useace>1</useace>
   <resultcolumns></resultcolumns>
   <template></template>
   <iscombinatortemplate></iscombinatortemplate>
   <allowmultiplestdins></allowmultiplestdins>
   <answer><![CDATA[def harf(not_):
    if not_ < 0 or not_ > 100:
        return 'Gecersiz'
    if not_ >= 90:
        return 'AA'
    if not_ >= 80:
        return 'BA'
    if not_ >= 70:
        return 'BB'
    if not_ >= 60:
        return 'CB'
    if not_ >= 50:
        return 'CC'
    return 'FF'
]]></answer>
   <validateonsave>1</validateonsave>
   <testsplitterre></testsplitterre>
   <language></language>
   <acelang></acelang>
   <sandbox></sandbox>
   <grader></grader>
   <cputimelimitsecs></cputimelimitsecs>
   <memlimitmb></memlimitmb>
   <sandboxparams></sandboxparams>
   <testcases>
   <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(harf(95))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[AA]]></text></expected>
    <extra><text></text></extra>
    <display><text>SHOW</text></display>
   </testcase>
   <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(harf(90))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[AA]]></text></expected>
    <extra><text></text></extra>
    <display><text>SHOW</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(harf(89))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[BA]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(harf(80))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[BA]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(harf(79))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[BB]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(harf(70))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[BB]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(harf(69))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[CB]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(harf(60))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[CB]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(harf(59))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[CC]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(harf(50))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[CC]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(harf(49))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[FF]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(harf(0))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[FF]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(harf(-1))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[Gecersiz]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(harf(101))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[Gecersiz]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(harf(100))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[AA]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   </testcases>
  </question>
  <question type="coderunner">
   <name><text>s3_asal - Asal mı? [Paket A]</text></name>
   <questiontext format="html"><text><![CDATA[<p><code>asal_mi(n)</code> fonksiyonu <code>n</code> asal ise <code>True</code>, değilse <code>False</code> döndürsün. (0, 1 ve negatifler asal değildir.)</p>]]></text></questiontext>
   <generalfeedback format="html"><text></text></generalfeedback>
   <defaultgrade>15</defaultgrade>
   <penalty>0.1</penalty>
   <hidden>0</hidden>
   <coderunnertype>python3</coderunnertype>
   <prototypetype>0</prototypetype>
   <allornothing>0</allornothing>
   <penaltyregime>10, 20, ...</penaltyregime>
   <precheck>0</precheck>
   <hidecheck>0</hidecheck>
   <showsource>0</showsource>
   <answerboxlines>18</answerboxlines>
   <answerboxcolumns>100</answerboxcolumns>
   <answerpreload></answerpreload>
   <globalextra></globalextra>
   <useace>1</useace>
   <resultcolumns></resultcolumns>
   <template></template>
   <iscombinatortemplate></iscombinatortemplate>
   <allowmultiplestdins></allowmultiplestdins>
   <answer><![CDATA[def asal_mi(n):
    if n < 2:
        return False
    i = 2
    while i * i <= n:
        if n % i == 0:
            return False
        i += 1
    return True
]]></answer>
   <validateonsave>1</validateonsave>
   <testsplitterre></testsplitterre>
   <language></language>
   <acelang></acelang>
   <sandbox></sandbox>
   <grader></grader>
   <cputimelimitsecs></cputimelimitsecs>
   <memlimitmb></memlimitmb>
   <sandboxparams></sandboxparams>
   <testcases>
   <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(asal_mi(2))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[True]]></text></expected>
    <extra><text></text></extra>
    <display><text>SHOW</text></display>
   </testcase>
   <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(asal_mi(3))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[True]]></text></expected>
    <extra><text></text></extra>
    <display><text>SHOW</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(asal_mi(4))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[False]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(asal_mi(1))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[False]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(asal_mi(0))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[False]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(asal_mi(-7))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[False]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(asal_mi(17))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[True]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(asal_mi(25))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[False]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(asal_mi(97))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[True]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(asal_mi(100))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[False]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(asal_mi(7919))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[True]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   </testcases>
  </question>
  <question type="coderunner">
   <name><text>s4_palindrom - Palindrom [Paket A]</text></name>
   <questiontext format="html"><text><![CDATA[<p><code>palindrom_mu(s)</code> fonksiyonu; büyük/küçük harf ve <b>boşlukları yok sayarak</b> metin palindrom ise <code>True</code>, değilse <code>False</code> döndürsün.</p>

<p>Örnek: <code>&quot;Kay ak&quot;</code> → <code>True</code>.</p>]]></text></questiontext>
   <generalfeedback format="html"><text></text></generalfeedback>
   <defaultgrade>15</defaultgrade>
   <penalty>0.1</penalty>
   <hidden>0</hidden>
   <coderunnertype>python3</coderunnertype>
   <prototypetype>0</prototypetype>
   <allornothing>0</allornothing>
   <penaltyregime>10, 20, ...</penaltyregime>
   <precheck>0</precheck>
   <hidecheck>0</hidecheck>
   <showsource>0</showsource>
   <answerboxlines>18</answerboxlines>
   <answerboxcolumns>100</answerboxcolumns>
   <answerpreload></answerpreload>
   <globalextra></globalextra>
   <useace>1</useace>
   <resultcolumns></resultcolumns>
   <template></template>
   <iscombinatortemplate></iscombinatortemplate>
   <allowmultiplestdins></allowmultiplestdins>
   <answer><![CDATA[def palindrom_mu(s):
    t = s.replace(' ', '').lower()
    return t == t[::-1]
]]></answer>
   <validateonsave>1</validateonsave>
   <testsplitterre></testsplitterre>
   <language></language>
   <acelang></acelang>
   <sandbox></sandbox>
   <grader></grader>
   <cputimelimitsecs></cputimelimitsecs>
   <memlimitmb></memlimitmb>
   <sandboxparams></sandboxparams>
   <testcases>
   <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(palindrom_mu("kayak"))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[True]]></text></expected>
    <extra><text></text></extra>
    <display><text>SHOW</text></display>
   </testcase>
   <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(palindrom_mu("Kay ak"))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[True]]></text></expected>
    <extra><text></text></extra>
    <display><text>SHOW</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(palindrom_mu("python"))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[False]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(palindrom_mu(""))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[True]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(palindrom_mu("A man a plan a canal Panama"))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[True]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(palindrom_mu("ab"))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[False]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[print(palindrom_mu("Aa"))]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[True]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   </testcases>
  </question>
  <question type="coderunner">
   <name><text>s5_frekans - Harf frekansı [Paket A]</text></name>
   <questiontext format="html"><text><![CDATA[<p>Bir satır metin okuyun. Harfleri (küçük harfe çevirerek, boşluk ve diğer karakterleri atlayarak) sayın. Her harfi <code>harf: adet</code> biçiminde, <b>önce adedi çok olan, eşitlikte alfabetik</b> sırayla ayrı satırlarda yazdırın.</p>]]></text></questiontext>
   <generalfeedback format="html"><text></text></generalfeedback>
   <defaultgrade>20</defaultgrade>
   <penalty>0.1</penalty>
   <hidden>0</hidden>
   <coderunnertype>python3</coderunnertype>
   <prototypetype>0</prototypetype>
   <allornothing>0</allornothing>
   <penaltyregime>10, 20, ...</penaltyregime>
   <precheck>0</precheck>
   <hidecheck>0</hidecheck>
   <showsource>0</showsource>
   <answerboxlines>18</answerboxlines>
   <answerboxcolumns>100</answerboxcolumns>
   <answerpreload></answerpreload>
   <globalextra></globalextra>
   <useace>1</useace>
   <resultcolumns></resultcolumns>
   <template></template>
   <iscombinatortemplate></iscombinatortemplate>
   <allowmultiplestdins></allowmultiplestdins>
   <answer><![CDATA[s = input().lower()
sayac = {}
for c in s:
    if c.isalpha():
        sayac[c] = sayac.get(c, 0) + 1
for harf, adet in sorted(sayac.items(), key=lambda x: (-x[1], x[0])):
    print(f'{harf}: {adet}')
]]></answer>
   <validateonsave>1</validateonsave>
   <testsplitterre></testsplitterre>
   <language></language>
   <acelang></acelang>
   <sandbox></sandbox>
   <grader></grader>
   <cputimelimitsecs></cputimelimitsecs>
   <memlimitmb></memlimitmb>
   <sandboxparams></sandboxparams>
   <testcases>
   <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[]]></text></testcode>
    <stdin><text><![CDATA[banana
]]></text></stdin>
    <expected><text><![CDATA[a: 3
n: 2
b: 1]]></text></expected>
    <extra><text></text></extra>
    <display><text>SHOW</text></display>
   </testcase>
   <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[]]></text></testcode>
    <stdin><text><![CDATA[Hello World
]]></text></stdin>
    <expected><text><![CDATA[l: 3
o: 2
d: 1
e: 1
h: 1
r: 1
w: 1]]></text></expected>
    <extra><text></text></extra>
    <display><text>SHOW</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[]]></text></testcode>
    <stdin><text><![CDATA[aabbcc
]]></text></stdin>
    <expected><text><![CDATA[a: 2
b: 2
c: 2]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[]]></text></testcode>
    <stdin><text><![CDATA[a
]]></text></stdin>
    <expected><text><![CDATA[a: 1]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[]]></text></testcode>
    <stdin><text><![CDATA[x y z x
]]></text></stdin>
    <expected><text><![CDATA[x: 2
y: 1
z: 1]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[]]></text></testcode>
    <stdin><text><![CDATA[123 !!
]]></text></stdin>
    <expected><text><![CDATA[]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[]]></text></testcode>
    <stdin><text><![CDATA[Python ve Java
]]></text></stdin>
    <expected><text><![CDATA[a: 2
v: 2
e: 1
h: 1
j: 1
n: 1
o: 1
p: 1
t: 1
y: 1]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   </testcases>
  </question>
  <question type="coderunner">
   <name><text>s6_banka - Banka hesabı (OOP) [Paket A]</text></name>
   <questiontext format="html"><text><![CDATA[<p><code>Hesap</code> sınıfını yazın:</p>

<ul>
<li><code>Hesap(sahip, bakiye=0)</code> — kurucu</li>
<li><code>yatir(miktar)</code> — miktar pozitifse bakiyeye ekler; değilse <code>ValueError</code> fırlatır</li>
<li><code>cek(miktar)</code> — yetersiz bakiyede <code>ValueError(&quot;Yetersiz bakiye&quot;)</code> fırlatır; aksi halde düşer</li>
<li><code>__str__</code> → <code>Sahip: ali, Bakiye: 100.00</code></li>
</ul>]]></text></questiontext>
   <generalfeedback format="html"><text></text></generalfeedback>
   <defaultgrade>25</defaultgrade>
   <penalty>0.1</penalty>
   <hidden>0</hidden>
   <coderunnertype>python3</coderunnertype>
   <prototypetype>0</prototypetype>
   <allornothing>0</allornothing>
   <penaltyregime>10, 20, ...</penaltyregime>
   <precheck>0</precheck>
   <hidecheck>0</hidecheck>
   <showsource>0</showsource>
   <answerboxlines>18</answerboxlines>
   <answerboxcolumns>100</answerboxcolumns>
   <answerpreload></answerpreload>
   <globalextra></globalextra>
   <useace>1</useace>
   <resultcolumns></resultcolumns>
   <template></template>
   <iscombinatortemplate></iscombinatortemplate>
   <allowmultiplestdins></allowmultiplestdins>
   <answer><![CDATA[class Hesap:
    def __init__(self, sahip, bakiye=0):
        self.sahip = sahip
        self.bakiye = bakiye

    def yatir(self, miktar):
        if miktar <= 0:
            raise ValueError('Gecersiz miktar')
        self.bakiye += miktar

    def cek(self, miktar):
        if miktar > self.bakiye:
            raise ValueError('Yetersiz bakiye')
        self.bakiye -= miktar

    def __str__(self):
        return f'Sahip: {self.sahip}, Bakiye: {self.bakiye:.2f}'
]]></answer>
   <validateonsave>1</validateonsave>
   <testsplitterre></testsplitterre>
   <language></language>
   <acelang></acelang>
   <sandbox></sandbox>
   <grader></grader>
   <cputimelimitsecs></cputimelimitsecs>
   <memlimitmb></memlimitmb>
   <sandboxparams></sandboxparams>
   <testcases>
   <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[h = Hesap('ali', 100)
print(h)]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[Sahip: ali, Bakiye: 100.00]]></text></expected>
    <extra><text></text></extra>
    <display><text>SHOW</text></display>
   </testcase>
   <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[h = Hesap('veli')
h.yatir(50)
h.yatir(25.5)
print(h)]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[Sahip: veli, Bakiye: 75.50]]></text></expected>
    <extra><text></text></extra>
    <display><text>SHOW</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[h = Hesap('ayse', 100)
h.cek(30)
print(h)]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[Sahip: ayse, Bakiye: 70.00]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[h = Hesap('x', 10)
try:
    h.cek(50)
except ValueError as e:
    print(e)
print(h)]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[Yetersiz bakiye
Sahip: x, Bakiye: 10.00]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[h = Hesap('y')
try:
    h.yatir(-5)
except ValueError:
    print('hata')
print(h.bakiye)]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[hata
0]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
    <testcode><text><![CDATA[h = Hesap('z', 20)
h.cek(20)
print(h)]]></text></testcode>
    <stdin><text><![CDATA[]]></text></stdin>
    <expected><text><![CDATA[Sahip: z, Bakiye: 0.00]]></text></expected>
    <extra><text></text></extra>
    <display><text>HIDE</text></display>
   </testcase>
   </testcases>
  </question>
</quiz>
